_Lambda School Data Science, Unit 2_

---

👇 **Do not change the code in this cell.** If you're working in Google Colab, you can run this cell to install `category_encoders` and `pdpbox`.

```python
%%capture
import sys

if 'google.colab' in sys.modules:
    !pip install category_encoders
    !pip install pdpbox
```

# Sprint Challenge: Predict Chicago Food Inspections 🍕

In this challenge, you'll use data from the [Chicago Department of Public Health](https://www.chicago.gov/city/en/depts/cdph/provdrs/healthy_restaurants/svcs/food-protection-services.html) to build a model to predict whether a food establishment passed inspection or not.

The purpose of this model is to help inspectors use their time more efficiently by identifying establishments that will likely fail inspection. In other words, this model should be able to predict whether an establishment will fail inspection *before* the inspector arrives at the establishment. 

# Directions

This notebook contains 12 tasks, which cover the material we've learned in this sprint. Here's a summary:

- **Task 1:** Importing data.
- **Task 2:** Identifying data leakage.
- **Task 3:** Writing a wrangle function.
- **Task 4:** Splitting data into a feature matrix and target vector.
- **Task 5:** Splitting data into training and validation sets.
- **Task 6:** Establishing baseline accuracy.
- **Task 7:** Building model with bagging predictor.
- **Task 8:** Building model with boosting predictor.
- **Task 9 (`stretch goal`):** Plotting ROC curves.
- **Task 10:** Generating classification report.
- **Task 11:** Calculating permutation importances.
- **Task 12 (`stretch goal`):** Creating PDP interaction plot.

For each task you should do the following:

- Read the task instructions.
- Write your code in the cell below the task. Delete the `raise NotImplementedError` before your start.
- Run the testing cell below the task. If you get an error, read the error message and re-evaluate your code.

**You should limit your code to the following libraries:**

- `category_encoders`
- `numpy`
- `matplotlib`
- `pandas`
- `pdpbox`
- `sklearn`
- `xgboost`

# I. Wrangle Data

**Task 1:** Change the code below to import your dataset. Be sure to examine the columns carefully and determine if one of them should be set as the index.

<!-- grade_id: cell-8b9246d8d97a80ff, points: 0, grade: false, solution: true -->
```python
'''T1. Import data file.'''
import pandas as pd
import numpy as np
url = 'https://drive.google.com/uc?export=download&id=1aUnQ4AJK4UtW8JL9zPyYUMtkjIgQpqKT'
df = pd.read_csv(url, parse_dates=['Inspection Date'], index_col='Inspection Date')
# YOUR CODE HERE

```

**Task 1 Test**

```python
df.head()
```

<!-- grade_id: cell-e9593d4f4ed7a9bb, points: 1, grade: true, solution: false -->
```python
'''T1 Test'''
assert isinstance(df, pd.DataFrame), 'Have you created a DataFrame named `df`?'
assert len(df) == 51916
```

**Task 2:** Given that this model is supposed to generate predictions *before* an inspection is conducted, identify the numerical feature that is an example of **data leakage.** Assign the column name to the variable `'leaky_col'`.

**Remember:** Leakage is when your feature matrix includes columns that will not be available to your model at the time it make predictions.



<!-- grade_id: cell-ef24afc9168ad64f, points: 0, grade: false, solution: true -->
```python
'''T2. Identify data leakage column.'''
leaky_col = 'Serious Violations Found'
# YOUR CODE HERE

```

**Task 2 Test**

<!-- grade_id: cell-8429f30efb2a7bf7, points: 1, grade: true, solution: false -->
```python
'''T2 Test'''
# This is a hidden test. 
# You'll see the result when you submit to Canvas.
```

Output:
```
'T2 Test'
```

**Task 3:** Add to the `wrangle` function below so that it does the following:

- Removes the "leaky" column.
- Removes high-cardinality categorical columns.
- Removes categorical columns that have only one category.
- Removes numerical columns that are unique identifiers for each observation, not features that would affect the target.

```python
df
```

Output:
```
Inspection ID  ... Serious Violations Found
Inspection Date                 ...                         
2017-09-15             2088270  ...                      2.0
2011-10-20              555268  ...                      0.0
2016-04-05             1751394  ...                      0.0
2016-04-29             1763905  ...                      0.0
2011-01-10              453326  ...                      0.0
...                        ...  ...                      ...
2013-10-29             1370030  ...                      NaN
2011-02-25              579387  ...                      NaN
2018-05-14             2169678  ...                      0.0
2015-11-30             1588969  ...                      NaN
2010-09-29              114560  ...                      0.0

[51916 rows x 17 columns]
```

<!-- grade_id: cell-d6fc5ee398afff4f, points: 0, grade: false, solution: true -->
```python
'''T3. Write wrangle function.'''
def wrangle(X):
  X.fillna(method='ffill', inplace=True)
  high_card_cols = [col for col in X.select_dtypes('object').columns
                      if X[col].nunique() > 500]
  X.drop(columns=high_card_cols, inplace=True)
  cols = [leaky_col, 'State','License #','Latitude','Longitude']
  X.drop(columns=cols, inplace=True)                   
  return df
# YOUR CODE HERE

df = wrangle(df)
```

**Task 3 Test**

<!-- grade_id: cell-49f495efb58bcd9f, points: 1, grade: true, solution: false -->
```python
'''T3 Test'''
assert df.select_dtypes('object').nunique().max() < 500, 'Have you dropped the high-cardinality columns?'
assert df.select_dtypes('object').nunique().min() > 1, 'Have you dropped the column with only one category?'
```

```python
df
```

Output:
```
Inspection ID             Risk  ...          Inspection Type  Fail
Inspection Date                                  ...                               
2017-09-15             2088270    Risk 1 (High)  ...                  Canvass     1
2011-10-20              555268     Risk 3 (Low)  ...  Complaint Re-Inspection     0
2016-04-05             1751394    Risk 1 (High)  ...    Canvass Re-Inspection     0
2016-04-29             1763905  Risk 2 (Medium)  ...                  Canvass     0
2011-01-10              453326    Risk 1 (High)  ...                  Canvass     0
...                        ...              ...  ...                      ...   ...
2013-10-29             1370030  Risk 2 (Medium)  ...                  Canvass     0
2011-02-25              579387  Risk 2 (Medium)  ...    License Re-Inspection     0
2018-05-14             2169678    Risk 1 (High)  ...    Canvass Re-Inspection     0
2015-11-30             1588969    Risk 1 (High)  ...    Canvass Re-Inspection     0
2010-09-2
```

# II. Split Data

**Task 4:** Split the DataFrame `df` into the feature matrix `X` and the target vector `y`. Your target is `'Fail'`.

<!-- grade_id: cell-b21b1c40f5478337, points: 0, grade: false, solution: true -->
```python
'''T4. Split feature matrix and target vector.'''
target = 'Fail'
# YOUR CODE HERE
y = df[target]
X = df.drop(target, axis = 1)
```

**Task 4 Test**

<!-- grade_id: cell-a1d912e28c9f7522, points: 1, grade: true, solution: false -->
```python
'''T4 Test'''
assert y.shape == (51916,), '`y` either has the wrong number of rows, or is two-dimentional.'
assert len(X) == 51916, '`X` has the wrong number of rows.'
```

**Task 5:** Split your dataset into training and validation sets.

- Your training set (`X_train`, `y_train`) should contain inspections conducted before 2017.
- Your validation set (`X_val`, `y_val`) should contain inspections conducted during or after 2017.

<!-- grade_id: cell-0bb47689fd4667ed, points: 0, grade: false, solution: true -->
```python
'''T5. Split dataset into training and validation sets.'''
# YOUR CODE HERE
mask = X.index < '2016-12-31'

X_train, y_train = X.loc[mask], y.loc[mask]
X_val, y_val = X.loc[~mask], y.loc[~mask]
```

**Task 5 Testing**

<!-- grade_id: cell-52cf3ef1934a4278, points: 1, grade: true, solution: false -->
```python
'''T5 Test'''
assert len(X_train) == len(y_train) == 41827, 'Your training set has the wrong number of observations.'
assert len(X_val) == len(y_val) == 10089, 'Your validation set has the wrong number of observations.'
```

# III. Establish Baseline

**Task 6:** Establish the baseline accuracy score for this classification problem using your training set. Save the score to the variable `baseline_acc`.

<!-- grade_id: cell-3d21cc97649be107, points: 0, grade: false, solution: true -->
```python
'''T6. Establish baseline accuracy.'''
# YOUR CODE HERE
baseline_acc = y_train.value_counts(normalize=True).max()
print('Baseline accuracy:', baseline_acc)
```

Output:
```
Baseline accuracy: 0.7506634470557295
```

**Task 6 Testing**

<!-- grade_id: cell-abdc4cbe95e9d1da, points: 1, grade: true, solution: false -->
```python
'''T6 Test'''
assert isinstance(baseline_acc, float), '`baseline_acc` should be type float. Have you defined the variable?'
assert 0.0 <= baseline_acc <= 1.0
```

# IV. Build Model

In this section, you want to answer the question: Which ensemble method performs better with this data — bagging or boosting?

**Task 7:** Build a model that includes a bagging predictor (`RandomForest`). Your predictor should be part of a pipeline named `model_bag` that includes any transformers that you think are necessary.

```python
df.isnull().sum()
```

Output:
```
Inspection ID       0
Risk               12
City               53
Zip                26
Inspection Type     1
Fail                0
dtype: int64
```

```python
from sklearn.pipeline import make_pipeline
from sklearn.ensemble import RandomForestClassifier
from sklearn.impute import SimpleImputer
from sklearn.preprocessing import OrdinalEncoder
```

<!-- grade_id: cell-889285d53fdbe282, points: 0, grade: false, solution: true -->
```python
'''T7. Build model with bagging predictor.'''
# YOUR CODE HERE
model_bag = make_pipeline(
    OrdinalEncoder(),
    SimpleImputer(),
    RandomForestClassifier()
)
model_bag.fit(X_train, y_train)
model_bag.fit(X_val,y_val)
```

Output:
```
Pipeline(memory=None,
         steps=[('ordinalencoder',
                 OrdinalEncoder(categories='auto',
                                dtype=<class 'numpy.float64'>)),
                ('simpleimputer',
                 SimpleImputer(add_indicator=False, copy=True, fill_value=None,
                               missing_values=nan, strategy='mean',
                               verbose=0)),
                ('randomforestclassifier',
                 RandomForestClassifier(bootstrap=True, ccp_alpha=0.0,
                                        class_weight=None, criterion='gini',
                                        max_depth=None, max_features='auto',
                                        max_leaf_nodes=None, max_samples=None,
                                        min_impurity_decrease=0.0,
                                        min_impurity_split=None,
                                        min_samples_leaf=1, min_samples_split=2,
                                        m
```

**Tast 7 Testing**

<!-- grade_id: cell-cddc5d7d2170877b, points: 1, grade: true, solution: false -->
```python
'''T7 Testing'''
from sklearn.pipeline import Pipeline
from sklearn.ensemble import RandomForestClassifier
assert isinstance(model_bag, Pipeline), '`model_bag` is the wrong data type. Have you assigned your pipeline to the correct variable name?'
assert isinstance(model_bag[-1], RandomForestClassifier), 'Your predictor should be a `RandomForestClassifier`.'
assert hasattr(model_bag[-1], 'feature_importances_'), 'Have you trained your model?'
```

**Task 8:** Build a model that includes a boosting predictor (`GradientBoostingClassifier` from `sklearn` or `XGBClassifier` from `xgboost`). Your predictor should be part of a pipeline named `model_boost` that includes any transformers that you think are necessary. 

```python
from xgboost import XGBClassifier
```

<!-- grade_id: cell-37f16b5811ae5223, points: 0, grade: false, solution: true -->
```python
'''T8. Build model with boosting predictor.'''
# YOUR CODE HERE
model_boost = make_pipeline(
    OrdinalEncoder(),
    SimpleImputer(),
    XGBClassifier()
)
model_boost.fit(X_train, y_train)
model_boost.fit(X_val, y_val)
```

Output:
```
Pipeline(memory=None,
         steps=[('ordinalencoder',
                 OrdinalEncoder(categories='auto',
                                dtype=<class 'numpy.float64'>)),
                ('simpleimputer',
                 SimpleImputer(add_indicator=False, copy=True, fill_value=None,
                               missing_values=nan, strategy='mean',
                               verbose=0)),
                ('xgbclassifier',
                 XGBClassifier(base_score=0.5, booster='gbtree',
                               colsample_bylevel=1, colsample_bynode=1,
                               colsample_bytree=1, gamma=0, learning_rate=0.1,
                               max_delta_step=0, max_depth=3,
                               min_child_weight=1, missing=None,
                               n_estimators=100, n_jobs=1, nthread=None,
                               objective='binary:logistic', random_state=0,
                               reg_alpha=0, reg_lambda=1, scale_pos_weight=
```

**Task 8 Testing**

<!-- grade_id: cell-90deb42a1c052402, points: 1, grade: true, solution: false -->
```python
'''T8 Testing'''
from xgboost import XGBClassifier
from sklearn.ensemble import GradientBoostingClassifier
assert isinstance(model_boost, Pipeline), '`model_boost` is the wrong data type. Have you assigned your pipeline to the correct variable name?'
assert any([isinstance(model_boost[-1], XGBClassifier),
            isinstance(model_boost[-1], GradientBoostingClassifier)]), 'Your predictor should be `XGBClassifier` or `GradientBoostingClassifier`.'
```

# V. Check Metrics

Here are the accuracy scores for your two models. Did you beat the baseline? Which of your two models appears to perform better on your validation set?

```python
df.info()
```

Output:
```
<class 'pandas.core.frame.DataFrame'>
DatetimeIndex: 51916 entries, 2017-09-15 to 2010-09-29
Data columns (total 7 columns):
 #   Column           Non-Null Count  Dtype  
---  ------           --------------  -----  
 0   Inspection ID    51916 non-null  int64  
 1   Facility Type    51916 non-null  object 
 2   Risk             51916 non-null  object 
 3   City             51916 non-null  object 
 4   Zip              51916 non-null  float64
 5   Inspection Type  51916 non-null  object 
 6   Fail             51916 non-null  int64  
dtypes: float64(1), int64(2), object(4)
memory usage: 3.2+ MB
```

```python
print('Bagging Model')
print('Training accuracy:', model_bag.score(X_train, y_train))
print('Validation accuracy:', model_bag.score(X_val, y_val))
print()
print('Boosting Model')
print('Training accuracy:', model_boost.score(X_train, y_train))
print('Validation accuracy:', model_boost.score(X_val, y_val))
```

Output:
```
Bagging Model
/usr/local/lib/python3.6/dist-packages/sklearn/preprocessing/_encoders.py in transform(self, X)
    645             Transformed input.
    646         """
--> 647         X_int, _ = self._transform(X)
    648         return X_int.astype(self.dtype, copy=False)
    649 

/usr/local/lib/python3.6/dist-packages/sklearn/preprocessing/_encoders.py in _transform(self, X, handle_unknown)
    122                     msg = ("Found unknown categories {0} in column {1}"
    123                            " during transform".format(diff, i))
--> 124                     raise ValueError(msg)
    125                 else:
    126                     # Set the problematic rows to an acceptable value and

ValueError: Found unknown categories [44250, 44251, 44254, 44258, 48212, 48213, 48216, 48219, 48220, 48222, 48224, 48225, 48226, 48229, 52235, 52237, 52238, 52239, 52240, 52247, 52249, 52250, 52256, 52257, 52270, 54212, 54218, 54219, 54220, 54224, 54236, 54248, 54253, 58249, 58250, 5825
```

**Task 9 (`stretch_goal`):** Plot the ROC-curve for both of your models (you can plot them one-at-a-time, side-by-side, or in the same plot).

<!-- grade_id: cell-769e4a780bb22283, points: 0, grade: false, solution: true -->
```python
'''T9. Plot ROC-curve.'''
# YOUR CODE HERE
raise NotImplementedError()
```

**Task 10:** Choose one of your models based on your validation accuracy score or your ROC curves. Then create a classification report for that model using your validation data. Save the text of the report to the variable name `model_cr`. 

<!-- grade_id: cell-49891c4ce9bf5f37, points: 0, grade: false, solution: true -->
```python
'''T10. Generate classification report for one model.'''
from sklearn.metrics import classification_report
# YOUR CODE HERE
raise NotImplementedError()
print(model_cr)
```

**Task 10 Testing**

<!-- grade_id: cell-94e04c938f3f5f84, points: 1, grade: true, solution: false -->
```python
assert isinstance(model_cr, str), '`model_cr` should be type `str`.'
assert all(term in model_cr for term in ['precision', 'recall', 'f1-score', 'support']), 'Is this a classification report?'
```

**Task 11:** Using your best model, create a DataFrame `permutation_importances` with the model's permutation importances based on your validation data.

- The index of the DataFrame should be your feature names.
- The first column should be the mean importance.
- The second column should be the importance standard deviation.

<!-- grade_id: cell-72936eec6980072b, points: 0, grade: false, solution: true -->
```python
'''Task 11. Create DataFrame of permutation importances.'''
# YOUR CODE HERE
raise NotImplementedError()
permutation_importances
```

**Task 11 Testing**

<!-- grade_id: cell-a4d8990e7070c2dd, points: 1, grade: true, solution: false -->
```python
'''Task 11 Test'''
assert isinstance(permutation_importances, pd.DataFrame), '`permutation_importances` should be type `DataFrame`.'
assert permutation_importances.shape == (7,2)
```

**Task 12 (`stretch goal`):** Using your best model, create a PDP interaction plot to examine how `'Latitude'` and `'Longitude'` inform predictions. Remember to user your validation data.

**Note:** Because of the way that `pdp_interact` works, it will throw an error if there are `NaN` values in your validation set. To avoid this problem, be sure to set `dataset` to `X_val.dropna()`.

<!-- grade_id: cell-224d3b408f9bdd88, points: 0, grade: false, solution: true -->
```python
'''T12. Create PDP interaction plot for "Latitude" and "Longitude".'''
features = ['Longitude', 'Latitude']
# YOUR CODE HERE
raise NotImplementedError()
```

What do you think? Is there a relationship between location and failing a food saftey inspection?